File,Vampire
halting2_FAIL.n3s,% (42519)input_syntax=auto not supported for standard input parsing, switching to tptp.% (42527)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1199ds/107Mi)% (42525)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1199ds/258102Mi)% (42525)Refutation not found, incomplete strategy% (42525)------------------------------% (42525)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42525)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42525)Termination reason: Refutation not found, incomplete strategy% (42525)Memory used [KB]: 1031% (42525)Time elapsed: 0.002 s% (42525)Instructions burned: 4 (million)% (42523)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1199ds/222662Mi)% (42524)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1199ds/257909Mi)% (42525)------------------------------% (42525)------------------------------% (42526)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1199ds/103Mi)% (42528)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1199ds/116Mi)% (42523)Refutation not found, incomplete strategy% (42523)------------------------------% (42523)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42523)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42523)Termination reason: Refutation not found, incomplete strategy% (42523)Memory used [KB]: 1042% (42523)Time elapsed: 0.002 s% (42523)Instructions burned: 4 (million)% (42528)Refutation not found, incomplete strategy% (42528)------------------------------% (42528)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42528)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42528)Termination reason: Refutation not found, incomplete strategy% (42528)Memory used [KB]: 1044% (42528)Time elapsed: 0.003 s% (42528)Instructions burned: 4 (million)% (42528)------------------------------% (42528)------------------------------% (42523)------------------------------% (42523)------------------------------% (42527)First to succeed.% (42524)Also succeeded, but the first one will report.% (42527)Solution written to "/tmp/vampire-proof-42519"% (42526)Also succeeded, but the first one will report.% (42527)Refutation found. Thanks to Tanya!% SZS status Theorem for % SZS output start Proof for fof(f84,plain,(  $false),  inference(avatar_sat_refutation,[],[f18,f39,f83])).fof(f83,plain,(  ~spl1_2),  inference(avatar_contradiction_clause,[],[f82])).fof(f82,plain,(  $false | ~spl1_2),  inference(subsumption_resolution,[],[f79,f37])).fof(f37,plain,(  triple(sK0(sK0(sK0(sK0('urn:example:T1')))),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(unit_resulting_resolution,[],[f26,f8])).fof(f8,plain,(  ( ! [X10] : (triple(sK0(X10),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X10,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(cnf_transformation,[],[f6])).fof(f6,plain,(  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (~triple(X9,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X8,'urn:example:child',X9) | ~triple(X7,'urn:example:child',X8) | ~triple(X6,'urn:example:child',X7) | ~triple(X4,'urn:example:child',X5) | ~triple(X3,'urn:example:child',X4) | ~triple(X2,'urn:example:child',X3) | ~triple(X1,'urn:example:child',X2) | ~triple(X0,'urn:example:child',X1) | ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & ! [X10] : (? [X11] : (triple(X10,'urn:example:child',X11) & triple(X11,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) | ~triple(X10,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(ennf_transformation,[],[f5])).fof(f5,plain,(  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ~(triple(X9,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') & triple(X8,'urn:example:child',X9) & triple(X7,'urn:example:child',X8) & triple(X6,'urn:example:child',X7) & triple(X4,'urn:example:child',X5) & triple(X3,'urn:example:child',X4) & triple(X2,'urn:example:child',X3) & triple(X1,'urn:example:child',X2) & triple(X0,'urn:example:child',X1) & triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & ! [X10] : ~(! [X11] : ~(triple(X10,'urn:example:child',X11) & triple(X11,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & triple(X10,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(rectify,[],[f1])).fof(f1,axiom,(  ! [X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ~(triple(X11,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') & triple(X10,'urn:example:child',X11) & triple(X9,'urn:example:child',X10) & triple(X8,'urn:example:child',X9) & triple(X6,'urn:example:child',X7) & triple(X5,'urn:example:child',X6) & triple(X4,'urn:example:child',X5) & triple(X3,'urn:example:child',X4) & triple(X2,'urn:example:child',X3) & triple(X2,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & ! [X0] : ~(! [X1] : ~(triple(X0,'urn:example:child',X1) & triple(X1,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  file(unknown,unknown)).fof(f26,plain,(  triple(sK0(sK0(sK0('urn:example:T1'))),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(unit_resulting_resolution,[],[f22,f8])).fof(f22,plain,(  triple(sK0(sK0('urn:example:T1')),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(unit_resulting_resolution,[],[f19,f8])).fof(f19,plain,(  triple(sK0('urn:example:T1'),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(unit_resulting_resolution,[],[f11,f8])).fof(f11,plain,(  triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(cnf_transformation,[],[f6])).fof(f79,plain,(  ~triple(sK0(sK0(sK0(sK0('urn:example:T1')))),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~spl1_2),  inference(unit_resulting_resolution,[],[f64,f9])).fof(f9,plain,(  ( ! [X10] : (triple(X10,'urn:example:child',sK0(X10)) | ~triple(X10,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(cnf_transformation,[],[f6])).fof(f64,plain,(  ( ! [X0] : (~triple(sK0(sK0(sK0(sK0('urn:example:T1')))),'urn:example:child',X0)) ) | ~spl1_2),  inference(unit_resulting_resolution,[],[f11,f20,f21,f25,f36,f17])).fof(f17,plain,(  ( ! [X2,X3,X0,X1,X4,X5] : (~triple(X4,'urn:example:child',X5) | ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X3,'urn:example:child',X4) | ~triple(X2,'urn:example:child',X3) | ~triple(X1,'urn:example:child',X2) | ~triple(X0,'urn:example:child',X1)) ) | ~spl1_2),  inference(avatar_component_clause,[],[f16])).fof(f16,plain,(  spl1_2 <=> ! [X2,X4,X0,X5,X3,X1] : (~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X4,'urn:example:child',X5) | ~triple(X3,'urn:example:child',X4) | ~triple(X2,'urn:example:child',X3) | ~triple(X1,'urn:example:child',X2) | ~triple(X0,'urn:example:child',X1))),  introduced(avatar_definition,[new_symbols(naming,[spl1_2])])).fof(f36,plain,(  triple(sK0(sK0(sK0('urn:example:T1'))),'urn:example:child',sK0(sK0(sK0(sK0('urn:example:T1')))))),  inference(unit_resulting_resolution,[],[f26,f9])).fof(f25,plain,(  triple(sK0(sK0('urn:example:T1')),'urn:example:child',sK0(sK0(sK0('urn:example:T1'))))),  inference(unit_resulting_resolution,[],[f22,f9])).fof(f21,plain,(  triple(sK0('urn:example:T1'),'urn:example:child',sK0(sK0('urn:example:T1')))),  inference(unit_resulting_resolution,[],[f19,f9])).fof(f20,plain,(  triple('urn:example:T1','urn:example:child',sK0('urn:example:T1'))),  inference(unit_resulting_resolution,[],[f11,f9])).fof(f39,plain,(  ~spl1_1),  inference(avatar_contradiction_clause,[],[f38])).fof(f38,plain,(  $false | ~spl1_1),  inference(subsumption_resolution,[],[f35,f25])).fof(f35,plain,(  ~triple(sK0(sK0('urn:example:T1')),'urn:example:child',sK0(sK0(sK0('urn:example:T1')))) | ~spl1_1),  inference(unit_resulting_resolution,[],[f21,f20,f26,f14])).fof(f14,plain,(  ( ! [X8,X6,X9,X7] : (~triple(X9,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X6,'urn:example:child',X7) | ~triple(X8,'urn:example:child',X9) | ~triple(X7,'urn:example:child',X8)) ) | ~spl1_1),  inference(avatar_component_clause,[],[f13])).fof(f13,plain,(  spl1_1 <=> ! [X9,X7,X6,X8] : (~triple(X6,'urn:example:child',X7) | ~triple(X9,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X8,'urn:example:child',X9) | ~triple(X7,'urn:example:child',X8))),  introduced(avatar_definition,[new_symbols(naming,[spl1_1])])).fof(f18,plain,(  spl1_1 | spl1_2),  inference(avatar_split_clause,[],[f10,f16,f13])).fof(f10,plain,(  ( ! [X2,X3,X0,X1,X8,X6,X9,X7,X4,X5] : (~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X0,'urn:example:child',X1) | ~triple(X1,'urn:example:child',X2) | ~triple(X2,'urn:example:child',X3) | ~triple(X3,'urn:example:child',X4) | ~triple(X4,'urn:example:child',X5) | ~triple(X6,'urn:example:child',X7) | ~triple(X7,'urn:example:child',X8) | ~triple(X8,'urn:example:child',X9) | ~triple(X9,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(cnf_transformation,[],[f6])).% SZS output end Proof for Consequence
negativeSurface_FAIL.n3s,% (42579)input_syntax=auto not supported for standard input parsing, switching to tptp.% (42584)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1199ds/103Mi)% (42582)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1199ds/257909Mi)% (42583)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1199ds/258102Mi)% (42581)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1199ds/222662Mi)% (42585)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1199ds/107Mi)% (42586)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1199ds/116Mi)% (42582)Also succeeded, but the first one will report.% (42583)Also succeeded, but the first one will report.% (42584)First to succeed.% (42585)Also succeeded, but the first one will report.% (42581)Also succeeded, but the first one will report.% (42584)Solution written to "/tmp/vampire-proof-42579"% (42586)Also succeeded, but the first one will report.% (42584)Refutation found. Thanks to Tanya!% SZS status ContradictoryAxioms for % SZS output start Proof for fof(f7,plain,(  $false),  inference(cnf_transformation,[],[f5])).fof(f5,plain,(  $false),  inference(true_and_false_elimination,[],[f1])).fof(f1,axiom,(  ~$true),  file(unknown,unknown)).% SZS output end Proof for Contradiction
contradiction_FAIL.n3s,% (42636)input_syntax=auto not supported for standard input parsing, switching to tptp.% (42641)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1199ds/257909Mi)% (42643)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1199ds/103Mi)% (42641)First to succeed.% (42643)Also succeeded, but the first one will report.% (42641)Solution written to "/tmp/vampire-proof-42636"% (42640)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1199ds/222662Mi)% (42642)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1199ds/258102Mi)% (42644)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1199ds/107Mi)% (42640)Refutation not found, incomplete strategy% (42640)------------------------------% (42640)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42640)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42640)Termination reason: Refutation not found, incomplete strategy% (42640)Memory used [KB]: 1038% (42640)Time elapsed: 0.003 s% (42640)Instructions burned: 4 (million)% (42642)Refutation not found, incomplete strategy% (42642)------------------------------% (42642)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42642)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42642)Termination reason: Refutation not found, incomplete strategy% (42642)Memory used [KB]: 1031% (42642)Time elapsed: 0.002 s% (42642)Instructions burned: 4 (million)% (42645)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1199ds/116Mi)% (42641)Refutation found. Thanks to Tanya!% SZS status ContradictoryAxioms for % SZS output start Proof for fof(f10,plain,(  $false),  inference(subsumption_resolution,[],[f8,f9])).fof(f9,plain,(  ( ! [X0] : (~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example.org:Cat')) )),  inference(cnf_transformation,[],[f6])).fof(f6,plain,(  (triple('urn:example.org:test','urn:example.org:is','"true"^^http://www.w3.org/2001/XMLSchema#boolean') | ~triple('urn:example.org:test','urn:example.org:is','"true"^^http://www.w3.org/2001/XMLSchema#boolean')) & ! [X0] : ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example.org:Cat') & ! [X1] : triple(X1,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example.org:Cat')),  inference(ennf_transformation,[],[f5])).fof(f5,plain,(  ~(~triple('urn:example.org:test','urn:example.org:is','"true"^^http://www.w3.org/2001/XMLSchema#boolean') & triple('urn:example.org:test','urn:example.org:is','"true"^^http://www.w3.org/2001/XMLSchema#boolean')) & ! [X0] : ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example.org:Cat') & ! [X1] : triple(X1,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example.org:Cat')),  inference(rectify,[],[f1])).fof(f1,axiom,(  ~(~triple('urn:example.org:test','urn:example.org:is','"true"^^http://www.w3.org/2001/XMLSchema#boolean') & triple('urn:example.org:test','urn:example.org:is','"true"^^http://www.w3.org/2001/XMLSchema#boolean')) & ! [X0] : ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example.org:Cat') & ! [X0] : triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example.org:Cat')),  file(unknown,unknown)).fof(f8,plain,(  ( ! [X1] : (triple(X1,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example.org:Cat')) )),  inference(cnf_transformation,[],[f6])).% SZS output end Proof for Contradiction
syllogism5_FAIL.n3s,% (42691)input_syntax=auto not supported for standard input parsing, switching to tptp.% (42695)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1199ds/258102Mi)% (42693)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1199ds/222662Mi)% (42694)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1199ds/257909Mi)% (42695)Refutation not found, incomplete strategy% (42695)------------------------------% (42695)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42696)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1199ds/103Mi)% (42695)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42695)Termination reason: Refutation not found, incomplete strategy% (42695)Memory used [KB]: 1031% (42695)Time elapsed: 0.003 s% (42695)Instructions burned: 4 (million)% (42697)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1199ds/107Mi)% (42698)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1199ds/116Mi)% (42693)Refutation not found, incomplete strategy% (42693)------------------------------% (42693)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42693)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42693)Termination reason: Refutation not found, incomplete strategy% (42693)Memory used [KB]: 1036% (42693)Time elapsed: 0.002 s% (42693)Instructions burned: 4 (million)% (42694)First to succeed.% (42695)------------------------------% (42695)------------------------------% (42696)Also succeeded, but the first one will report.% (42697)Also succeeded, but the first one will report.% (42698)Refutation not found, incomplete strategy% (42698)------------------------------% (42698)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42698)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42698)Termination reason: Refutation not found, incomplete strategy% (42698)Memory used [KB]: 1039% (42698)Time elapsed: 0.003 s% (42698)Instructions burned: 4 (million)% (42693)------------------------------% (42693)------------------------------% (42694)Solution written to "/tmp/vampire-proof-42691"% (42698)------------------------------% (42698)------------------------------% (42694)Refutation found. Thanks to Tanya!% SZS status ContradictoryAxioms for % SZS output start Proof for fof(f13,plain,(  $false),  inference(subsumption_resolution,[],[f12,f9])).fof(f9,plain,(  triple('urn:example:Alice','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Person')),  inference(cnf_transformation,[],[f5])).fof(f5,plain,(  ! [X0] : ((triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:IceSkater') & triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Cyclist')) | ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Person')) & ~triple('urn:example:Alice','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:IceSkater') & triple('urn:example:Alice','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Cyclist') & triple('urn:example:Alice','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Person')),  inference(ennf_transformation,[],[f1])).fof(f1,axiom,(  ! [X0] : ~(~(triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:IceSkater') & triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Cyclist')) & triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Person')) & ~triple('urn:example:Alice','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:IceSkater') & triple('urn:example:Alice','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Cyclist') & triple('urn:example:Alice','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Person')),  file(unknown,unknown)).fof(f12,plain,(  ~triple('urn:example:Alice','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Person')),  inference(resolution,[],[f8,f11])).fof(f11,plain,(  ~triple('urn:example:Alice','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:IceSkater')),  inference(cnf_transformation,[],[f5])).fof(f8,plain,(  ( ! [X0] : (triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:IceSkater') | ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Person')) )),  inference(cnf_transformation,[],[f5])).% SZS output end Proof for Contradiction
mona2_FAIL.n3s,% (42752)input_syntax=auto not supported for standard input parsing, switching to tptp.% (42758)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1199ds/103Mi)% (42760)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1199ds/116Mi)% (42756)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1199ds/257909Mi)% (42757)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1199ds/258102Mi)% (42755)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1199ds/222662Mi)% (42759)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1199ds/107Mi)% (42760)Refutation not found, incomplete strategy% (42760)------------------------------% (42760)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42760)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42760)Termination reason: Refutation not found, incomplete strategy% (42760)Memory used [KB]: 1038% (42760)Time elapsed: 0.002 s% (42760)Instructions burned: 4 (million)% (42756)Also succeeded, but the first one will report.% (42758)First to succeed.% (42757)Refutation not found, incomplete strategy% (42757)------------------------------% (42757)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42757)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42757)Termination reason: Refutation not found, incomplete strategy% (42757)Memory used [KB]: 1032% (42757)Time elapsed: 0.002 s% (42757)Instructions burned: 4 (million)% (42755)Refutation not found, incomplete strategy% (42755)------------------------------% (42755)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42755)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42755)Termination reason: Refutation not found, incomplete strategy% (42755)Memory used [KB]: 1037% (42755)Time elapsed: 0.002 s% (42760)------------------------------% (42760)------------------------------% (42755)Instructions burned: 4 (million)% (42757)------------------------------% (42757)------------------------------% (42759)Also succeeded, but the first one will report.% (42758)Solution written to "/tmp/vampire-proof-42752"% (42755)------------------------------% (42755)------------------------------% (42758)Refutation found. Thanks to Tanya!% SZS status Theorem for % SZS output start Proof for fof(f48,plain,(  $false),  inference(avatar_sat_refutation,[],[f29,f34,f39,f44,f45,f47])).fof(f47,plain,(  ~spl0_3),  inference(avatar_contradiction_clause,[],[f46])).fof(f46,plain,(  $false | ~spl0_3),  inference(resolution,[],[f26,f14])).fof(f14,plain,(  ~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works')),  inference(cnf_transformation,[],[f5])).fof(f5,plain,(  ~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works') & (triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works') | triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works2') | ~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works1')) & (triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works') | (~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works2') & ~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works1'))) & (triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works1') | (~triple('http://example.org/ns#Mona','http://example.org/ns#has','http://example.org/ns#acrophobia') & triple('http://example.org/ns#Mona','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','http://example.org/ns#Cow'))) & (triple('http://example.org/ns#Mona','http://example.org/ns#has','http://example.org/ns#acrophobia') | triple('http://example.org/ns#Mona','http://example.org/ns#color','http://example.org/ns#purple') | ~triple('http://example.org/ns#Mona','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','http://example.org/ns#Cow')) & (triple('http://example.org/ns#Mona','http://example.org/ns#has','http://example.org/ns#acrophobia') | ~triple('http://example.org/ns#Mona','http://example.org/ns#color','http://example.org/ns#purple') | ~triple('http://example.org/ns#Mona','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','http://example.org/ns#Cow'))),  inference(ennf_transformation,[],[f1])).fof(f1,axiom,(  ~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works') & ~(~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works') & ~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works2') & triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works1')) & ~(~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works') & ~(~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works2') & ~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works1'))) & ~(~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works1') & ~(~triple('http://example.org/ns#Mona','http://example.org/ns#has','http://example.org/ns#acrophobia') & triple('http://example.org/ns#Mona','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','http://example.org/ns#Cow'))) & ~(~triple('http://example.org/ns#Mona','http://example.org/ns#has','http://example.org/ns#acrophobia') & ~triple('http://example.org/ns#Mona','http://example.org/ns#color','http://example.org/ns#purple') & triple('http://example.org/ns#Mona','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','http://example.org/ns#Cow')) & ~(~triple('http://example.org/ns#Mona','http://example.org/ns#has','http://example.org/ns#acrophobia') & triple('http://example.org/ns#Mona','http://example.org/ns#color','http://example.org/ns#purple') & triple('http://example.org/ns#Mona','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','http://example.org/ns#Cow'))),  file(unknown,unknown)).fof(f26,plain,(  triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works') | ~spl0_3),  inference(avatar_component_clause,[],[f24])).fof(f24,plain,(  spl0_3 <=> triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works')),  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).fof(f45,plain,(  ~spl0_5 | ~spl0_6 | spl0_4),  inference(avatar_split_clause,[],[f7,f31,f41,f36])).fof(f36,plain,(  spl0_5 <=> triple('http://example.org/ns#Mona','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','http://example.org/ns#Cow')),  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).fof(f41,plain,(  spl0_6 <=> triple('http://example.org/ns#Mona','http://example.org/ns#color','http://example.org/ns#purple')),  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).fof(f31,plain,(  spl0_4 <=> triple('http://example.org/ns#Mona','http://example.org/ns#has','http://example.org/ns#acrophobia')),  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).fof(f7,plain,(  triple('http://example.org/ns#Mona','http://example.org/ns#has','http://example.org/ns#acrophobia') | ~triple('http://example.org/ns#Mona','http://example.org/ns#color','http://example.org/ns#purple') | ~triple('http://example.org/ns#Mona','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','http://example.org/ns#Cow')),  inference(cnf_transformation,[],[f5])).fof(f44,plain,(  ~spl0_5 | spl0_6 | spl0_4),  inference(avatar_split_clause,[],[f8,f31,f41,f36])).fof(f8,plain,(  triple('http://example.org/ns#Mona','http://example.org/ns#has','http://example.org/ns#acrophobia') | triple('http://example.org/ns#Mona','http://example.org/ns#color','http://example.org/ns#purple') | ~triple('http://example.org/ns#Mona','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','http://example.org/ns#Cow')),  inference(cnf_transformation,[],[f5])).fof(f39,plain,(  spl0_5 | spl0_1),  inference(avatar_split_clause,[],[f9,f16,f36])).fof(f16,plain,(  spl0_1 <=> triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works1')),  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).fof(f9,plain,(  triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works1') | triple('http://example.org/ns#Mona','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','http://example.org/ns#Cow')),  inference(cnf_transformation,[],[f5])).fof(f34,plain,(  ~spl0_4 | spl0_1),  inference(avatar_split_clause,[],[f10,f16,f31])).fof(f10,plain,(  triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works1') | ~triple('http://example.org/ns#Mona','http://example.org/ns#has','http://example.org/ns#acrophobia')),  inference(cnf_transformation,[],[f5])).fof(f29,plain,(  ~spl0_1 | spl0_3),  inference(avatar_split_clause,[],[f11,f24,f16])).fof(f11,plain,(  triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works') | ~triple('http://example.org/ns#this','http://example.org/ns#test','http://example.org/ns#works1')),  inference(cnf_transformation,[],[f5])).% SZS output end Proof for Consequence
contradiction4_FAIL.n3s,% (42811)input_syntax=auto not supported for standard input parsing, switching to tptp.% (42814)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1199ds/257909Mi)% (42813)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1199ds/222662Mi)% (42815)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1199ds/258102Mi)% (42816)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1199ds/103Mi)% (42815)Refutation not found, incomplete strategy% (42815)------------------------------% (42815)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42815)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42815)Termination reason: Refutation not found, incomplete strategy% (42815)Memory used [KB]: 1030% (42815)Time elapsed: 0.003 s% (42815)Instructions burned: 4 (million)% (42813)Refutation not found, incomplete strategy% (42813)------------------------------% (42813)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42813)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42813)Termination reason: Refutation not found, incomplete strategy% (42818)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1199ds/116Mi)% (42813)Memory used [KB]: 1034% (42813)Time elapsed: 0.002 s% (42813)Instructions burned: 4 (million)% (42814)First to succeed.% (42816)Also succeeded, but the first one will report.% (42817)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1199ds/107Mi)% (42814)Solution written to "/tmp/vampire-proof-42811"% (42815)------------------------------% (42815)------------------------------% (42818)Refutation not found, incomplete strategy% (42818)------------------------------% (42818)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42813)------------------------------% (42813)------------------------------% (42818)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42818)Termination reason: Refutation not found, incomplete strategy% (42818)Memory used [KB]: 1034% (42818)Time elapsed: 0.003 s% (42818)Instructions burned: 4 (million)% (42817)Also succeeded, but the first one will report.% (42818)------------------------------% (42818)------------------------------% (42814)Refutation found. Thanks to Tanya!% SZS status ContradictoryAxioms for % SZS output start Proof for fof(f10,plain,(  $false),  inference(subsumption_resolution,[],[f8,f9])).fof(f9,plain,(  ~triple('http://example.org/ns#Alice','http://example.org/ns#hates','http://example.org/ns#Quiche')),  inference(cnf_transformation,[],[f6])).fof(f6,plain,(  ~triple('http://example.org/ns#Alice','http://example.org/ns#hates','http://example.org/ns#Quiche') & triple('http://example.org/ns#Alice','http://example.org/ns#hates','http://example.org/ns#Quiche')),  inference(flattening,[],[f5])).fof(f5,plain,(  ~triple('http://example.org/ns#Alice','http://example.org/ns#hates','http://example.org/ns#Quiche') & ~~triple('http://example.org/ns#Alice','http://example.org/ns#hates','http://example.org/ns#Quiche')),  inference(rectify,[],[f1])).fof(f1,axiom,(  ! [X0] : ~triple('http://example.org/ns#Alice','http://example.org/ns#hates','http://example.org/ns#Quiche') & ~! [X0,X1,X2] : ~triple('http://example.org/ns#Alice','http://example.org/ns#hates','http://example.org/ns#Quiche')),  file(unknown,unknown)).fof(f8,plain,(  triple('http://example.org/ns#Alice','http://example.org/ns#hates','http://example.org/ns#Quiche')),  inference(cnf_transformation,[],[f6])).% SZS output end Proof for Contradiction
russell_FAIL.n3s,% (42872)input_syntax=auto not supported for standard input parsing, switching to tptp.% (42875)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1199ds/257909Mi)% (42876)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1199ds/258102Mi)% (42877)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1199ds/103Mi)% (42878)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1199ds/107Mi)% (42879)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1199ds/116Mi)% (42876)Refutation not found, incomplete strategy% (42876)------------------------------% (42876)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42876)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42876)Termination reason: Refutation not found, incomplete strategy% (42876)Memory used [KB]: 1030% (42876)Time elapsed: 0.002 s% (42876)Instructions burned: 4 (million)% (42877)Also succeeded, but the first one will report.% (42875)First to succeed.% (42874)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1199ds/222662Mi)% (42879)Refutation not found, incomplete strategy% (42879)------------------------------% (42879)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42879)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42879)Termination reason: Refutation not found, incomplete strategy% (42879)Memory used [KB]: 1038% (42879)Time elapsed: 0.002 s% (42879)Instructions burned: 4 (million)% (42878)Also succeeded, but the first one will report.% (42876)------------------------------% (42876)------------------------------% (42875)Solution written to "/tmp/vampire-proof-42872"% (42874)Refutation not found, incomplete strategy% (42874)------------------------------% (42874)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42874)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42874)Termination reason: Refutation not found, incomplete strategy% (42874)Memory used [KB]: 1037% (42874)Time elapsed: 0.003 s% (42874)Instructions burned: 4 (million)% (42879)------------------------------% (42879)------------------------------% (42874)------------------------------% (42874)------------------------------% (42875)Refutation found. Thanks to Tanya!% SZS status ContradictoryAxioms for % SZS output start Proof for fof(f12,plain,(  $false),  inference(subsumption_resolution,[],[f11,f8])).fof(f8,plain,(  ( ! [X0] : (~triple('http://example.org/ns#b','http://example.org/ns#shaves',X0) | ~triple(X0,'http://example.org/ns#shaves',X0)) )),  inference(cnf_transformation,[],[f6])).fof(f6,plain,(  ! [X0] : (~triple(X0,'http://example.org/ns#shaves',X0) | ~triple('http://example.org/ns#b','http://example.org/ns#shaves',X0)) & ! [X1] : (triple('http://example.org/ns#b','http://example.org/ns#shaves',X1) | triple(X1,'http://example.org/ns#shaves',X1))),  inference(ennf_transformation,[],[f5])).fof(f5,plain,(  ! [X0] : ~(triple(X0,'http://example.org/ns#shaves',X0) & triple('http://example.org/ns#b','http://example.org/ns#shaves',X0)) & ! [X1] : ~(~triple('http://example.org/ns#b','http://example.org/ns#shaves',X1) & ~triple(X1,'http://example.org/ns#shaves',X1))),  inference(rectify,[],[f1])).fof(f1,axiom,(  ! [X0] : ~(triple(X0,'http://example.org/ns#shaves',X0) & triple('http://example.org/ns#b','http://example.org/ns#shaves',X0)) & ! [X0] : ~(~triple('http://example.org/ns#b','http://example.org/ns#shaves',X0) & ~triple(X0,'http://example.org/ns#shaves',X0))),  file(unknown,unknown)).fof(f11,plain,(  triple('http://example.org/ns#b','http://example.org/ns#shaves','http://example.org/ns#b')),  inference(factoring,[],[f9])).fof(f9,plain,(  ( ! [X1] : (triple('http://example.org/ns#b','http://example.org/ns#shaves',X1) | triple(X1,'http://example.org/ns#shaves',X1)) )),  inference(cnf_transformation,[],[f6])).% SZS output end Proof for Contradiction
universal_statement2_FAIL.n3s,% (42933)input_syntax=auto not supported for standard input parsing, switching to tptp.% (42936)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1199ds/257909Mi)% (42937)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1199ds/258102Mi)% (42938)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1199ds/103Mi)% (42939)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1199ds/107Mi)% (42940)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1199ds/116Mi)% (42937)Refutation not found, incomplete strategy% (42937)------------------------------% (42937)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42937)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42937)Termination reason: Refutation not found, incomplete strategy% (42937)Memory used [KB]: 1031% (42937)Time elapsed: 0.002 s% (42937)Instructions burned: 4 (million)% (42936)First to succeed.% (42938)Also succeeded, but the first one will report.% (42939)Also succeeded, but the first one will report.% (42937)------------------------------% (42937)------------------------------% (42935)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1199ds/222662Mi)% (42940)Refutation not found, incomplete strategy% (42940)------------------------------% (42940)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42940)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42940)Termination reason: Refutation not found, incomplete strategy% (42940)Memory used [KB]: 1039% (42940)Time elapsed: 0.003 s% (42940)Instructions burned: 4 (million)% (42936)Solution written to "/tmp/vampire-proof-42933"% (42935)Refutation not found, incomplete strategy% (42935)------------------------------% (42935)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42935)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42935)Termination reason: Refutation not found, incomplete strategy% (42935)Memory used [KB]: 1035% (42935)Time elapsed: 0.003 s% (42935)Instructions burned: 4 (million)% (42935)------------------------------% (42935)------------------------------% (42936)Refutation found. Thanks to Tanya!% SZS status Theorem for % SZS output start Proof for fof(f30,plain,(  $false),  inference(avatar_sat_refutation,[],[f20,f25,f27,f29])).fof(f29,plain,(  ~spl1_3),  inference(avatar_contradiction_clause,[],[f28])).fof(f28,plain,(  $false | ~spl1_3),  inference(resolution,[],[f24,f8])).fof(f8,plain,(  triple('http://example.org/ns#This','http://example.org/ns#is','http://example.org/ns#working')),  inference(cnf_transformation,[],[f5])).fof(f5,plain,(  ~triple('http://example.org/ns#This','http://example.org/ns#is','http://example.org/ns#workingToo') & ! [X0,X1] : (triple(X1,'http://example.org/ns#is','http://example.org/ns#workingToo') | ~triple(X0,'http://example.org/ns#is','http://example.org/ns#working')) & triple('http://example.org/ns#This','http://example.org/ns#is','http://example.org/ns#working')),  inference(ennf_transformation,[],[f1])).fof(f1,axiom,(  ~triple('http://example.org/ns#This','http://example.org/ns#is','http://example.org/ns#workingToo') & ! [X0,X1] : ~(~triple(X1,'http://example.org/ns#is','http://example.org/ns#workingToo') & triple(X0,'http://example.org/ns#is','http://example.org/ns#working')) & triple('http://example.org/ns#This','http://example.org/ns#is','http://example.org/ns#working')),  file(unknown,unknown)).fof(f24,plain,(  ( ! [X0] : (~triple(X0,'http://example.org/ns#is','http://example.org/ns#working')) ) | ~spl1_3),  inference(avatar_component_clause,[],[f23])).fof(f23,plain,(  spl1_3 <=> ! [X0] : ~triple(X0,'http://example.org/ns#is','http://example.org/ns#working')),  introduced(avatar_definition,[new_symbols(naming,[spl1_3])])).fof(f27,plain,(  ~spl1_2),  inference(avatar_contradiction_clause,[],[f26])).fof(f26,plain,(  $false | ~spl1_2),  inference(resolution,[],[f19,f9])).fof(f9,plain,(  ~triple('http://example.org/ns#This','http://example.org/ns#is','http://example.org/ns#workingToo')),  inference(cnf_transformation,[],[f5])).fof(f19,plain,(  ( ! [X1] : (triple(X1,'http://example.org/ns#is','http://example.org/ns#workingToo')) ) | ~spl1_2),  inference(avatar_component_clause,[],[f18])).fof(f18,plain,(  spl1_2 <=> ! [X1] : triple(X1,'http://example.org/ns#is','http://example.org/ns#workingToo')),  introduced(avatar_definition,[new_symbols(naming,[spl1_2])])).fof(f25,plain,(  ~spl1_1 | spl1_3),  inference(avatar_split_clause,[],[f13,f23,f15])).fof(f15,plain,(  spl1_1 <=> sP0),  introduced(avatar_definition,[new_symbols(naming,[spl1_1])])).fof(f13,plain,(  ( ! [X0] : (~triple(X0,'http://example.org/ns#is','http://example.org/ns#working') | ~sP0) )),  inference(consistent_polarity_flipping,[],[f10])).fof(f10,plain,(  ( ! [X0] : (~triple(X0,'http://example.org/ns#is','http://example.org/ns#working') | sP0) )),  inference(cnf_transformation,[],[f10_D])).fof(f10_D,plain,(  ( ! [X0] : ~triple(X0,'http://example.org/ns#is','http://example.org/ns#working') ) <=> ~sP0),  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP0])])).fof(f20,plain,(  spl1_1 | spl1_2),  inference(avatar_split_clause,[],[f12,f18,f15])).fof(f12,plain,(  ( ! [X1] : (triple(X1,'http://example.org/ns#is','http://example.org/ns#workingToo') | sP0) )),  inference(consistent_polarity_flipping,[],[f11])).fof(f11,plain,(  ( ! [X1] : (triple(X1,'http://example.org/ns#is','http://example.org/ns#workingToo') | ~sP0) )),  inference(general_splitting,[],[f7,f10_D])).fof(f7,plain,(  ( ! [X0,X1] : (~triple(X0,'http://example.org/ns#is','http://example.org/ns#working') | triple(X1,'http://example.org/ns#is','http://example.org/ns#workingToo')) )),  inference(cnf_transformation,[],[f5])).% SZS output end Proof for Consequence
TheoLog24-Uebung-10_FAIL.n3s,% (42992)input_syntax=auto not supported for standard input parsing, switching to tptp.% (42998)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1199ds/258102Mi)% (42998)Refutation not found, incomplete strategy% (42998)------------------------------% (42998)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42998)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42998)Termination reason: Refutation not found, incomplete strategy% (42998)Memory used [KB]: 1031% (42998)Time elapsed: 0.002 s% (42998)Instructions burned: 4 (million)% (42999)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1199ds/103Mi)% (42996)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1199ds/222662Mi)% (42998)------------------------------% (42998)------------------------------% (42997)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1199ds/257909Mi)% (42996)Refutation not found, incomplete strategy% (42996)------------------------------% (42996)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (42996)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (42996)Termination reason: Refutation not found, incomplete strategy% (42996)Memory used [KB]: 1039% (42996)Time elapsed: 0.003 s% (42996)Instructions burned: 4 (million)% (42999)First to succeed.% (42997)Also succeeded, but the first one will report.% (43000)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1199ds/107Mi)% (42996)------------------------------% (42996)------------------------------% (43001)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1199ds/116Mi)% (42999)Solution written to "/tmp/vampire-proof-42992"% (43001)Refutation not found, incomplete strategy% (43001)------------------------------% (43001)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (43001)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (43001)Termination reason: Refutation not found, incomplete strategy% (43000)Also succeeded, but the first one will report.% (43001)Memory used [KB]: 1043% (43001)Time elapsed: 0.002 s% (43001)Instructions burned: 4 (million)% (43001)------------------------------% (43001)------------------------------% (42999)Refutation found. Thanks to Tanya!% SZS status Theorem for % SZS output start Proof for fof(f36,plain,(  $false),  inference(avatar_sat_refutation,[],[f28,f30,f35])).fof(f35,plain,(  ~spl2_1),  inference(avatar_contradiction_clause,[],[f34])).fof(f34,plain,(  $false | ~spl2_1),  inference(resolution,[],[f33,f16])).fof(f16,plain,(  ~triple(sK0,'http://example.org/ns#is','http://example.org/ns#happy')),  inference(cnf_transformation,[],[f9])).fof(f9,plain,(  (~triple(sK0,'http://example.org/ns#is','http://example.org/ns#happy') & triple(sK0,'http://example.org/ns#is','http://example.org/ns#green')) & ! [X1] : (triple(X1,'http://example.org/ns#is','http://example.org/ns#green') | ! [X2] : (~triple(X2,'http://example.org/ns#is','http://example.org/ns#green') | ~triple(X2,'http://example.org/ns#hasChild',X1))) & ! [X3] : (triple(X3,'http://example.org/ns#can','http://example.org/ns#fly') | ~triple(X3,'http://example.org/ns#is','http://example.org/ns#green')) & ! [X4] : (triple(X4,'http://example.org/ns#is','http://example.org/ns#happy') | (~triple(sK1(X4),'http://example.org/ns#can','http://example.org/ns#fly') & triple(X4,'http://example.org/ns#hasChild',sK1(X4))))),  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f6,f8,f7])).fof(f7,plain,(  ? [X0] : (~triple(X0,'http://example.org/ns#is','http://example.org/ns#happy') & triple(X0,'http://example.org/ns#is','http://example.org/ns#green')) => (~triple(sK0,'http://example.org/ns#is','http://example.org/ns#happy') & triple(sK0,'http://example.org/ns#is','http://example.org/ns#green'))),  introduced(choice_axiom,[])).fof(f8,plain,(  ! [X4] : (? [X5] : (~triple(X5,'http://example.org/ns#can','http://example.org/ns#fly') & triple(X4,'http://example.org/ns#hasChild',X5)) => (~triple(sK1(X4),'http://example.org/ns#can','http://example.org/ns#fly') & triple(X4,'http://example.org/ns#hasChild',sK1(X4))))),  introduced(choice_axiom,[])).fof(f6,plain,(  ? [X0] : (~triple(X0,'http://example.org/ns#is','http://example.org/ns#happy') & triple(X0,'http://example.org/ns#is','http://example.org/ns#green')) & ! [X1] : (triple(X1,'http://example.org/ns#is','http://example.org/ns#green') | ! [X2] : (~triple(X2,'http://example.org/ns#is','http://example.org/ns#green') | ~triple(X2,'http://example.org/ns#hasChild',X1))) & ! [X3] : (triple(X3,'http://example.org/ns#can','http://example.org/ns#fly') | ~triple(X3,'http://example.org/ns#is','http://example.org/ns#green')) & ! [X4] : (triple(X4,'http://example.org/ns#is','http://example.org/ns#happy') | ? [X5] : (~triple(X5,'http://example.org/ns#can','http://example.org/ns#fly') & triple(X4,'http://example.org/ns#hasChild',X5)))),  inference(ennf_transformation,[],[f5])).fof(f5,plain,(  ~! [X0] : ~(~triple(X0,'http://example.org/ns#is','http://example.org/ns#happy') & triple(X0,'http://example.org/ns#is','http://example.org/ns#green')) & ! [X1] : ~(~triple(X1,'http://example.org/ns#is','http://example.org/ns#green') & ~! [X2] : ~(triple(X2,'http://example.org/ns#is','http://example.org/ns#green') & triple(X2,'http://example.org/ns#hasChild',X1))) & ! [X3] : ~(~triple(X3,'http://example.org/ns#can','http://example.org/ns#fly') & triple(X3,'http://example.org/ns#is','http://example.org/ns#green')) & ! [X4] : ~(~triple(X4,'http://example.org/ns#is','http://example.org/ns#happy') & ! [X5] : ~(~triple(X5,'http://example.org/ns#can','http://example.org/ns#fly') & triple(X4,'http://example.org/ns#hasChild',X5)))),  inference(rectify,[],[f1])).fof(f1,axiom,(  ~! [X0] : ~(~triple(X0,'http://example.org/ns#is','http://example.org/ns#happy') & triple(X0,'http://example.org/ns#is','http://example.org/ns#green')) & ! [X0] : ~(~triple(X0,'http://example.org/ns#is','http://example.org/ns#green') & ~! [X1] : ~(triple(X1,'http://example.org/ns#is','http://example.org/ns#green') & triple(X1,'http://example.org/ns#hasChild',X0))) & ! [X0] : ~(~triple(X0,'http://example.org/ns#can','http://example.org/ns#fly') & triple(X0,'http://example.org/ns#is','http://example.org/ns#green')) & ! [X0] : ~(~triple(X0,'http://example.org/ns#is','http://example.org/ns#happy') & ! [X1] : ~(~triple(X1,'http://example.org/ns#can','http://example.org/ns#fly') & triple(X0,'http://example.org/ns#hasChild',X1)))),  file(unknown,unknown)).fof(f33,plain,(  triple(sK0,'http://example.org/ns#is','http://example.org/ns#happy') | ~spl2_1),  inference(resolution,[],[f32,f12])).fof(f12,plain,(  ( ! [X4] : (~triple(sK1(X4),'http://example.org/ns#can','http://example.org/ns#fly') | triple(X4,'http://example.org/ns#is','http://example.org/ns#happy')) )),  inference(cnf_transformation,[],[f9])).fof(f32,plain,(  triple(sK1(sK0),'http://example.org/ns#can','http://example.org/ns#fly') | ~spl2_1),  inference(resolution,[],[f23,f13])).fof(f13,plain,(  ( ! [X3] : (~triple(X3,'http://example.org/ns#is','http://example.org/ns#green') | triple(X3,'http://example.org/ns#can','http://example.org/ns#fly')) )),  inference(cnf_transformation,[],[f9])).fof(f23,plain,(  triple(sK1(sK0),'http://example.org/ns#is','http://example.org/ns#green') | ~spl2_1),  inference(avatar_component_clause,[],[f21])).fof(f21,plain,(  spl2_1 <=> triple(sK1(sK0),'http://example.org/ns#is','http://example.org/ns#green')),  introduced(avatar_definition,[new_symbols(naming,[spl2_1])])).fof(f30,plain,(  spl2_2),  inference(avatar_contradiction_clause,[],[f29])).fof(f29,plain,(  $false | spl2_2),  inference(resolution,[],[f27,f15])).fof(f15,plain,(  triple(sK0,'http://example.org/ns#is','http://example.org/ns#green')),  inference(cnf_transformation,[],[f9])).fof(f27,plain,(  ~triple(sK0,'http://example.org/ns#is','http://example.org/ns#green') | spl2_2),  inference(avatar_component_clause,[],[f25])).fof(f25,plain,(  spl2_2 <=> triple(sK0,'http://example.org/ns#is','http://example.org/ns#green')),  introduced(avatar_definition,[new_symbols(naming,[spl2_2])])).fof(f28,plain,(  spl2_1 | ~spl2_2),  inference(avatar_split_clause,[],[f19,f25,f21])).fof(f19,plain,(  ~triple(sK0,'http://example.org/ns#is','http://example.org/ns#green') | triple(sK1(sK0),'http://example.org/ns#is','http://example.org/ns#green')),  inference(resolution,[],[f18,f14])).fof(f14,plain,(  ( ! [X2,X1] : (~triple(X2,'http://example.org/ns#hasChild',X1) | ~triple(X2,'http://example.org/ns#is','http://example.org/ns#green') | triple(X1,'http://example.org/ns#is','http://example.org/ns#green')) )),  inference(cnf_transformation,[],[f9])).fof(f18,plain,(  triple(sK0,'http://example.org/ns#hasChild',sK1(sK0))),  inference(resolution,[],[f11,f16])).fof(f11,plain,(  ( ! [X4] : (triple(X4,'http://example.org/ns#is','http://example.org/ns#happy') | triple(X4,'http://example.org/ns#hasChild',sK1(X4))) )),  inference(cnf_transformation,[],[f9])).% SZS output end Proof for Consequence
halting3_FAIL.n3s,% (43050)input_syntax=auto not supported for standard input parsing, switching to tptp.% (43053)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1200ds/257909Mi)% (43055)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1200ds/103Mi)% (43052)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1200ds/222662Mi)% (43056)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1200ds/107Mi)% (43054)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1200ds/258102Mi)% (43057)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1200ds/116Mi)% (43054)Refutation not found, incomplete strategy% (43054)------------------------------% (43054)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (43054)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (43054)Termination reason: Refutation not found, incomplete strategy% (43054)Memory used [KB]: 1031% (43054)Time elapsed: 0.002 s% (43054)Instructions burned: 4 (million)% (43052)Refutation not found, incomplete strategy% (43052)------------------------------% (43052)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (43052)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (43052)Termination reason: Refutation not found, incomplete strategy% (43052)Memory used [KB]: 1040% (43052)Time elapsed: 0.002 s% (43052)Instructions burned: 4 (million)% (43057)Refutation not found, incomplete strategy% (43057)------------------------------% (43057)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (43057)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (43057)Termination reason: Refutation not found, incomplete strategy% (43057)Memory used [KB]: 1042% (43057)Time elapsed: 0.003 s% (43057)Instructions burned: 4 (million)% (43052)------------------------------% (43052)------------------------------% (43054)------------------------------% (43054)------------------------------% (43057)------------------------------% (43057)------------------------------% (43053)Also succeeded, but the first one will report.% (43055)First to succeed.% (43055)Solution written to "/tmp/vampire-proof-43050"% (43056)Also succeeded, but the first one will report.% (43055)Refutation found. Thanks to Tanya!% SZS status Theorem for % SZS output start Proof for fof(f131,plain,(  $false),  inference(avatar_sat_refutation,[],[f34,f130])).fof(f130,plain,(  ~spl1_1),  inference(avatar_split_clause,[],[f128,f25])).fof(f25,plain,(  spl1_1 <=> triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  introduced(avatar_definition,[new_symbols(naming,[spl1_1])])).fof(f128,plain,(  ~triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(resolution,[],[f119,f14])).fof(f14,plain,(  triple('urn:example:T1','urn:example:child',sK0('urn:example:T1'))),  inference(resolution,[],[f12,f10])).fof(f10,plain,(  triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(cnf_transformation,[],[f8])).fof(f8,plain,(  ! [X0,X1,X2,X3,X4,X5] : (~triple(X5,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X4,'urn:example:child',X5) | ~triple(X3,'urn:example:child',X4) | ~triple(X2,'urn:example:child',X3) | ~triple(X1,'urn:example:child',X2) | ~triple(X0,'urn:example:child',X1) | ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & ! [X6] : ((triple(X6,'urn:example:child',sK0(X6)) & triple(sK0(X6),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) | ~triple(X6,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f6,f7])).fof(f7,plain,(  ! [X6] : (? [X7] : (triple(X6,'urn:example:child',X7) & triple(X7,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) => (triple(X6,'urn:example:child',sK0(X6)) & triple(sK0(X6),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')))),  introduced(choice_axiom,[])).fof(f6,plain,(  ! [X0,X1,X2,X3,X4,X5] : (~triple(X5,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X4,'urn:example:child',X5) | ~triple(X3,'urn:example:child',X4) | ~triple(X2,'urn:example:child',X3) | ~triple(X1,'urn:example:child',X2) | ~triple(X0,'urn:example:child',X1) | ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & ! [X6] : (? [X7] : (triple(X6,'urn:example:child',X7) & triple(X7,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) | ~triple(X6,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(ennf_transformation,[],[f5])).fof(f5,plain,(  ! [X0,X1,X2,X3,X4,X5] : ~(triple(X5,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') & triple(X4,'urn:example:child',X5) & triple(X3,'urn:example:child',X4) & triple(X2,'urn:example:child',X3) & triple(X1,'urn:example:child',X2) & triple(X0,'urn:example:child',X1) & triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & ! [X6] : ~(! [X7] : ~(triple(X6,'urn:example:child',X7) & triple(X7,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & triple(X6,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  inference(rectify,[],[f1])).fof(f1,axiom,(  ! [X2,X3,X4,X5,X6,X7] : ~(triple(X7,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') & triple(X6,'urn:example:child',X7) & triple(X5,'urn:example:child',X6) & triple(X4,'urn:example:child',X5) & triple(X3,'urn:example:child',X4) & triple(X2,'urn:example:child',X3) & triple(X2,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & ! [X0] : ~(! [X1] : ~(triple(X0,'urn:example:child',X1) & triple(X1,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) & triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')),  file(unknown,unknown)).fof(f12,plain,(  ( ! [X6] : (~triple(X6,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | triple(X6,'urn:example:child',sK0(X6))) )),  inference(cnf_transformation,[],[f8])).fof(f119,plain,(  ( ! [X0] : (~triple('urn:example:T1','urn:example:child',sK0(X0)) | ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(duplicate_literal_removal,[],[f118])).fof(f118,plain,(  ( ! [X0] : (~triple('urn:example:T1','urn:example:child',sK0(X0)) | ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(resolution,[],[f111,f11])).fof(f11,plain,(  ( ! [X6] : (triple(sK0(X6),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X6,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(cnf_transformation,[],[f8])).fof(f111,plain,(  ( ! [X0] : (~triple(sK0(X0),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple('urn:example:T1','urn:example:child',sK0(X0)) | ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(resolution,[],[f101,f15])).fof(f15,plain,(  ( ! [X0] : (triple(sK0(X0),'urn:example:child',sK0(sK0(X0))) | ~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(resolution,[],[f12,f11])).fof(f101,plain,(  ( ! [X0,X1] : (~triple(X0,'urn:example:child',sK0(X1)) | ~triple('urn:example:T1','urn:example:child',X0) | ~triple(X1,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(duplicate_literal_removal,[],[f100])).fof(f100,plain,(  ( ! [X0,X1] : (~triple('urn:example:T1','urn:example:child',X0) | ~triple(X0,'urn:example:child',sK0(X1)) | ~triple(X1,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X1,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(resolution,[],[f87,f11])).fof(f87,plain,(  ( ! [X0,X1] : (~triple(sK0(X1),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple('urn:example:T1','urn:example:child',X0) | ~triple(X0,'urn:example:child',sK0(X1)) | ~triple(X1,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(resolution,[],[f78,f15])).fof(f78,plain,(  ( ! [X2,X0,X1] : (~triple(X1,'urn:example:child',sK0(X2)) | ~triple(X0,'urn:example:child',X1) | ~triple('urn:example:T1','urn:example:child',X0) | ~triple(X2,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(duplicate_literal_removal,[],[f77])).fof(f77,plain,(  ( ! [X2,X0,X1] : (~triple(X0,'urn:example:child',X1) | ~triple(X1,'urn:example:child',sK0(X2)) | ~triple('urn:example:T1','urn:example:child',X0) | ~triple(X2,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X2,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(resolution,[],[f58,f11])).fof(f58,plain,(  ( ! [X2,X0,X1] : (~triple(sK0(X2),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X0,'urn:example:child',X1) | ~triple(X1,'urn:example:child',sK0(X2)) | ~triple('urn:example:T1','urn:example:child',X0) | ~triple(X2,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(resolution,[],[f49,f15])).fof(f49,plain,(  ( ! [X2,X3,X0,X1] : (~triple(X2,'urn:example:child',sK0(X3)) | ~triple('urn:example:T1','urn:example:child',X0) | ~triple(X0,'urn:example:child',X1) | ~triple(X1,'urn:example:child',X2) | ~triple(X3,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(duplicate_literal_removal,[],[f48])).fof(f48,plain,(  ( ! [X2,X3,X0,X1] : (~triple(X0,'urn:example:child',X1) | ~triple('urn:example:T1','urn:example:child',X0) | ~triple(X2,'urn:example:child',sK0(X3)) | ~triple(X1,'urn:example:child',X2) | ~triple(X3,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X3,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(resolution,[],[f23,f11])).fof(f23,plain,(  ( ! [X2,X3,X0,X1] : (~triple(sK0(X3),'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X2,'urn:example:child',X0) | ~triple('urn:example:T1','urn:example:child',X2) | ~triple(X1,'urn:example:child',sK0(X3)) | ~triple(X0,'urn:example:child',X1) | ~triple(X3,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(resolution,[],[f19,f15])).fof(f19,plain,(  ( ! [X2,X3,X0,X1,X4] : (~triple(X1,'urn:example:child',sK0(X4)) | ~triple(X2,'urn:example:child',X0) | ~triple(X3,'urn:example:child',X2) | ~triple('urn:example:T1','urn:example:child',X3) | ~triple(X0,'urn:example:child',X1) | ~triple(X4,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(resolution,[],[f16,f11])).fof(f16,plain,(  ( ! [X2,X3,X0,X1,X4] : (~triple(X1,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X2,'urn:example:child',X0) | ~triple(X3,'urn:example:child',X2) | ~triple(X4,'urn:example:child',X3) | ~triple('urn:example:T1','urn:example:child',X4) | ~triple(X0,'urn:example:child',X1)) )),  inference(resolution,[],[f13,f10])).fof(f13,plain,(  ( ! [X2,X3,X0,X1,X4,X5] : (~triple(X0,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | ~triple(X4,'urn:example:child',X5) | ~triple(X3,'urn:example:child',X4) | ~triple(X2,'urn:example:child',X3) | ~triple(X1,'urn:example:child',X2) | ~triple(X0,'urn:example:child',X1) | ~triple(X5,'http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test')) )),  inference(cnf_transformation,[],[f8])).fof(f34,plain,(  spl1_1),  inference(avatar_contradiction_clause,[],[f33])).fof(f33,plain,(  $false | spl1_1),  inference(resolution,[],[f27,f10])).fof(f27,plain,(  ~triple('urn:example:T1','http://www.w3.org/1999/02/22-rdf-syntax-ns#type','urn:example:Test') | spl1_1),  inference(avatar_component_clause,[],[f25])).% SZS output end Proof for Consequence
contradiction2_FAIL.n3s,% (43110)input_syntax=auto not supported for standard input parsing, switching to tptp.% (43114)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1199ds/258102Mi)% (43112)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1199ds/222662Mi)% (43113)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1199ds/257909Mi)% (43114)Refutation not found, incomplete strategy% (43114)------------------------------% (43114)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (43114)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (43114)Termination reason: Refutation not found, incomplete strategy% (43115)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1199ds/103Mi)% (43114)Memory used [KB]: 1031% (43114)Time elapsed: 0.002 s% (43114)Instructions burned: 4 (million)% (43112)Refutation not found, incomplete strategy% (43112)------------------------------% (43112)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (43112)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (43112)Termination reason: Refutation not found, incomplete strategy% (43112)Memory used [KB]: 1036% (43112)Time elapsed: 0.003 s% (43112)Instructions burned: 4 (million)% (43117)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1199ds/116Mi)% (43114)------------------------------% (43114)------------------------------% (43115)Also succeeded, but the first one will report.% (43113)First to succeed.% (43112)------------------------------% (43112)------------------------------% (43117)Refutation not found, incomplete strategy% (43117)------------------------------% (43117)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (43117)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (43117)Termination reason: Refutation not found, incomplete strategy% (43117)Memory used [KB]: 1037% (43117)Time elapsed: 0.002 s% (43117)Instructions burned: 4 (million)% (43113)Solution written to "/tmp/vampire-proof-43110"% (43117)------------------------------% (43117)------------------------------% (43116)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1199ds/107Mi)% (43116)Also succeeded, but the first one will report.% (43113)Refutation found. Thanks to Tanya!% SZS status ContradictoryAxioms for % SZS output start Proof for fof(f10,plain,(  $false),  inference(resolution,[],[f9,f7])).fof(f7,plain,(  ( ! [X0] : (~triple('urn:example:Alice','urn:example:loves',X0)) )),  inference(cnf_transformation,[],[f5])).fof(f5,plain,(  (triple('urn:example:test','urn:example:is','"true"^^http://www.w3.org/2001/XMLSchema#boolean') | ~triple('urn:example:test','urn:example:is','"true"^^http://www.w3.org/2001/XMLSchema#boolean')) & (triple('urn:example:Alice','urn:example:loves','urn:example:Wine') | ~triple('urn:example:Alice','urn:example:likes','urn:example:Beer')) & ! [X0] : (~triple('urn:example:Alice','urn:example:loves',X0) | ~triple('urn:example:Alice','urn:example:loves',X0)) & triple('urn:example:Alice','urn:example:loves','urn:example:Wine')),  inference(ennf_transformation,[],[f1])).fof(f1,axiom,(  ~(~triple('urn:example:test','urn:example:is','"true"^^http://www.w3.org/2001/XMLSchema#boolean') & triple('urn:example:test','urn:example:is','"true"^^http://www.w3.org/2001/XMLSchema#boolean')) & ~(~triple('urn:example:Alice','urn:example:loves','urn:example:Wine') & triple('urn:example:Alice','urn:example:likes','urn:example:Beer')) & ! [X0] : ~(triple('urn:example:Alice','urn:example:loves',X0) & triple('urn:example:Alice','urn:example:loves',X0)) & triple('urn:example:Alice','urn:example:loves','urn:example:Wine')),  file(unknown,unknown)).fof(f9,plain,(  triple('urn:example:Alice','urn:example:loves','urn:example:Wine')),  inference(cnf_transformation,[],[f5])).% SZS output end Proof for Contradiction
russell2_FAIL.n3s,% (43169)input_syntax=auto not supported for standard input parsing, switching to tptp.% (43172)lrs+1011_4:1_sil=256000:rp=on:newcnf=on:i=257909:aac=none:gsp=on_0 on  for (1199ds/257909Mi)% (43173)dis+1002_1:1_tgt=full:sos=on:rp=on:sac=on:i=258102:ss=axioms:sd=3:cond=fast:add=off:abs=on:fde=none:sil=256000_0 on  for (1199ds/258102Mi)% (43174)lrs+21_8:1_to=lpo:sil=2000:sp=frequency:spb=units:s2a=on:s2pl=no:i=103:sd=2:ss=included:fsr=off:fs=off_0 on  for (1199ds/103Mi)% (43171)lrs+21_1:32_anc=all:to=lpo:sil=256000:plsq=on:plsqr=32,1:sp=occurrence:sos=on:plsql=on:sac=on:newcnf=on:i=222662:add=off:fsr=off:rawr=on_0 on  for (1199ds/222662Mi)% (43176)lrs+10_8:1_to=lpo:drc=encompass:sil=4000:sos=on:urr=on:newcnf=on:i=116:sd=2:nm=2:ss=axioms:sgt=32:sup=off:bd=off_0 on  for (1199ds/116Mi)% (43173)Refutation not found, incomplete strategy% (43173)------------------------------% (43173)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (43173)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (43173)Termination reason: Refutation not found, incomplete strategy% (43173)Memory used [KB]: 1030% (43173)Time elapsed: 0.003 s% (43173)Instructions burned: 4 (million)% (43175)lrs+1011_4:1_to=lpo:drc=off:sil=8000:sp=frequency:abs=on:urr=on:lsd=10:nwc=5.0:s2agt=4:newcnf=on:st=5.0:s2a=on:i=107:ss=axioms:aac=none:br=off:bd=preordered_0 on  for (1199ds/107Mi)% (43172)First to succeed.% (43174)Also succeeded, but the first one will report.% (43176)Refutation not found, incomplete strategy% (43176)------------------------------% (43176)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (43171)Refutation not found, incomplete strategy% (43171)------------------------------% (43171)Version: Vampire 4.9 (commit 5ad494e78 on 2024-06-14 14:05:27 +0100)% (43171)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (43171)Termination reason: Refutation not found, incomplete strategy% (43176)Linked with Z3 4.12.3.0 79bbbf76d0c123481c8ca05cd3a98939270074d3 z3-4.8.4-7980-g79bbbf76d% (43171)Memory used [KB]: 1037% (43176)Termination reason: Refutation not found, incomplete strategy% (43171)Time elapsed: 0.002 s% (43176)Memory used [KB]: 1038% (43171)Instructions burned: 4 (million)% (43176)Time elapsed: 0.003 s% (43176)Instructions burned: 4 (million)% (43173)------------------------------% (43173)------------------------------% (43172)Solution written to "/tmp/vampire-proof-43169"% (43175)Also succeeded, but the first one will report.% (43171)------------------------------% (43171)------------------------------% (43176)------------------------------% (43176)------------------------------% (43172)Refutation found. Thanks to Tanya!% SZS status ContradictoryAxioms for % SZS output start Proof for fof(f12,plain,(  $false),  inference(subsumption_resolution,[],[f11,f8])).fof(f8,plain,(  ( ! [X0] : (~triple('http://example.org/ns#b','http://example.org/ns#shaves',X0) | ~triple(X0,'http://example.org/ns#shaves',X0)) )),  inference(cnf_transformation,[],[f6])).fof(f6,plain,(  ! [X0] : (~triple(X0,'http://example.org/ns#shaves',X0) | ~triple('http://example.org/ns#b','http://example.org/ns#shaves',X0)) & ! [X1] : (triple('http://example.org/ns#b','http://example.org/ns#shaves',X1) | triple(X1,'http://example.org/ns#shaves',X1))),  inference(ennf_transformation,[],[f5])).fof(f5,plain,(  ! [X0] : ~(triple(X0,'http://example.org/ns#shaves',X0) & triple('http://example.org/ns#b','http://example.org/ns#shaves',X0)) & ! [X1] : ~(~triple('http://example.org/ns#b','http://example.org/ns#shaves',X1) & ~triple(X1,'http://example.org/ns#shaves',X1))),  inference(rectify,[],[f1])).fof(f1,axiom,(  ! [X0] : ~(triple(X0,'http://example.org/ns#shaves',X0) & triple('http://example.org/ns#b','http://example.org/ns#shaves',X0)) & ! [X0] : ~(~triple('http://example.org/ns#b','http://example.org/ns#shaves',X0) & ~triple(X0,'http://example.org/ns#shaves',X0))),  file(unknown,unknown)).fof(f11,plain,(  triple('http://example.org/ns#b','http://example.org/ns#shaves','http://example.org/ns#b')),  inference(factoring,[],[f9])).fof(f9,plain,(  ( ! [X1] : (triple('http://example.org/ns#b','http://example.org/ns#shaves',X1) | triple(X1,'http://example.org/ns#shaves',X1)) )),  inference(cnf_transformation,[],[f6])).% SZS output end Proof for Contradiction
